# Assignment 02

for the course [Introduction to Data Science and Programming](https://learnit.itu.dk/course/view.php?id=3022199) (Autumn 2023) at ITU

**Due: September 25, 23:59**

Before you turn this assignment in, make sure everything runs as expected. (Go to `Restart and run all` - should not throw any assertion errors.) 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and ITU email address below:

```python
NAME = ""
ITU_MAIL = "" # xxxx@itu.dk
```

***
**Before starting the tasks, import the modules you will need by running the cell below**

<!-- grade_id: cell-425f86bd60427ef2, points: 0, grade: false, solution: true -->
```python
# import math and random modules here
# YOUR CODE HERE
raise NotImplementedError()
```

***
**Mandatory task 1: At the restaurant** 

Define the function `calculate_bill` that computes the bill in a restaurant. The function takes 3 input parameters:
* `food` required paramter with a numeric value
* `drinks` required paramter with a numeric value
* `tip_percent` optional parameter with a default value of `0`

The function returns the total price on the restaurant bill, which includes the price of items (food+drinks), plus the tip (in **percent** of the price of all items).

If any of the provided values is non-numeric, your function should return the string `"error: input must be numeric"`.


<!-- grade_id: cell-6005ea839fca1142, points: 0, grade: false, solution: true -->
```python
# Solve task 1
def calculate_bill(food, drinks, tip_percent=0):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-33da5047a8925f34, points: 1, grade: true, solution: false -->
```python
# Check if you solved task 1 correctly

# if no tip is provided, the function should return just the sum of food & drinks:
assert calculate_bill(10,20)==30
assert calculate_bill(500,17.5)==517.5

# if a tip is provided (in PERCENT), the functon should return
# the sum of food & drinks PLUS the tip (as percent of food+drinks):
assert calculate_bill(40,20,10)==66
assert calculate_bill(100,60,20)==192

# if any of the arguments provided is of non-numeric type, 
# the function should throw the error:
assert calculate_bill(8,2,"none")=="error: input must be numeric"
assert calculate_bill(10,"3",5)=="error: input must be numeric"
assert calculate_bill(True,2,3)=="error: input must be numeric"

```

***
**Mandatory task 2: Your contact list**

Define the function `create_contact` that takes 4 parameters of personal data:
* `number` (has to be an integer, e.g. `"4530405099"`)
* `twitter` (string, which must start with an `@`, e.g. `@mszell`)
* `first_name` (string)
* `last_name` (string, optional, with the default value `"unknown"`)
and returns a dictionary which contains the parameters as keys, and the arguments as values. If no `last_name` is provided, the dictionary should have no `last_name` key.

If the function is called with any input argument that doesn't match the requirements (for example, if the `twitter` argument doesn't start with an `@`, or if the `number` argument contains letters), the function should return the string `"error: revise input"`.

<!-- grade_id: cell-3db6f0c321e415b6, points: 0, grade: false, solution: true -->
```python
# Solve task 2
def create_contact(number, twitter, first_name, last_name = "unknown"):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-64fd73b9b6bb7ec8, points: 1, grade: true, solution: false -->
```python
# Check if you solved task 2 correctly

# if all arguments are provided correctly, function returns the expected dictionary
# (the \ sign just means: "continue reading code in the next line")

assert create_contact(1, "@user", "user", "userovich") == \
    {
        "number": 1, 
        "twitter": "@user", 
        "first_name": "user", 
        "last_name": "userovich"
    }

assert create_contact(2211, "@rihanna", "Ri", "Hanna") == \
    {
        'number': 2211,
        'twitter': '@rihanna',
        'first_name': 'Ri',
        'last_name': 'Hanna'
    }

# if all arguments are provided correctly, and no last name is provided, function
# returns the expected dictionary (without the last name):
assert create_contact(2211, "@rihanna", "Ri") == \
    {
        'number': 2211,
        'twitter': '@rihanna',
        'first_name': 'Ri'
    }

# if some of the arguments is provided incorrectly, function returns the error:

# wrong number format
assert create_contact("+45204020", "@rihanna", "Ri") == "error: revise input"

# wrong twitter format
assert create_contact(45204020, "rihanna", "Ri") == "error: revise input"

# wrong name format (first or last - should both be str)
assert create_contact(45204020, "@rihanna", 123) == "error: revise input"
assert create_contact(45204020, "@rihanna", "Ri", 456) == "error: revise input"

```

***
**Mandatory task 3: Binomial coefficient**

Define the function `compute_binomial_coefficient` that takes the two non-negative (`>0`) integers `n` and `k` as input, and returns the value of their [binomial coefficient](https://en.wikipedia.org/wiki/Binomial_coefficient). Note that the binomial coefficient is only defined for `n>=k`.

The formula for the binomial coefficient is:

$\frac{n!}{k!*(n-k)!}$

Your function should return the string `"error: revise input"` if:
* any of the two input arguments is anything else than a non-negative integer;
* if `n<k`.

For this task, you are *allowed* to use any of the functions from the `math` module; but you don't have to.

<!-- grade_id: cell-557fca302a3191f5, points: 0, grade: false, solution: true -->
```python
# Solve task 3

# the math module was imported by running the first code cell in this notebook
def compute_binomial_coefficient(n, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-da256a0f8398f6b6, points: 0, grade: true, solution: false -->
```python
# Check if you solved task 3 correctly

# we get the expected values in cases where n and k are provided correctly
assert compute_binomial_coefficient(4,2)==6
assert compute_binomial_coefficient(3,1)==3
assert compute_binomial_coefficient(10,7)==120
assert compute_binomial_coefficient(10,3)==120
assert compute_binomial_coefficient(30,5)==142506

# we get the expected error message if n or k are of wrong type or <0, or if k>n:
assert compute_binomial_coefficient(10,20)=='error: revise input'
assert compute_binomial_coefficient(-5,10)=='error: revise input'
assert compute_binomial_coefficient(10,-5)=='error: revise input'
assert compute_binomial_coefficient(10,"4")=='error: revise input'
assert compute_binomial_coefficient(True,1)=='error: revise input'
```

***
**Mandatory task 4: Tossing coins**

Define the function `coin_tosses(n)` that takes an integer number `n>=1` as input; simulates `n` coin tosses, each represented by `H` or `T` (heads or tails; use `choice()` from the `random` module for this); and returns a string of length `n` that contains the results from all the coin tosses.

For this task, you don't have to implement any error catching (it is fine if your `coin_tosses` throws an error if anything other than an integer is provided as argument for `n`).



<!-- grade_id: cell-1643205e4cc9d6df, points: 0, grade: false, solution: true -->
```python
# Solve task 4

# the random module was imported by running the first code cell in this notebook
# so you can use the choice() function from the random module like so: random.choice()

def coin_tosses(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-87bad5d33886c213, points: 1, grade: true, solution: false -->
```python
# Try to check if you solved task 4 correctly 
# (but don't worry if the check fails)

random.seed(42) # don't change this line
assert coin_tosses(1)=='H'
assert coin_tosses(2)=='HT'
assert coin_tosses(5)=='HHHHH'
assert coin_tosses(10)=='THHHHHHHTH'

# NOTE: IF THIS CELL PRODUCES AN ERROR, DON'T WORRY,
# AS LONG AS YOUR FUNCTION PRODUCES THE EXPECTED OUTPUT 
# (strings of length n with a random combination of H and T)
```

***
**Mandatory task 5: Function parameters and arguments**

**In your own words**, explain the difference between function parameters and function arguments; and the difference between required and optional parameters. You can use your function `calculate_bill()` from Task 1 as example. 

**150 words max**

<!-- grade_id: cell-757386a6d7441f4b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE