# Assignment 04

for the course [Introduction to Data Science and Programming](https://learnit.itu.dk/course/view.php?id=3022199) (Autumn 2023) at ITU

**Due: October 9, 23:59**

Before you turn this assignment in, make sure everything runs as expected. (Go to `Restart and run all` - should not throw any assertion errors.) 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and ITU email address below:

```python
NAME = ""
ITU_MAIL = "" # xxxx@itu.dk
```

***
**Before starting the tasks, import the modules you will need by running the cell below**

<!-- grade_id: cell-4ff855f02bc93d54, points: 0, grade: false, solution: true -->
```python
# Use this cell to import all the packages needed for this exercise (os, pandas, requests) 
import os # etc.
# YOUR CODE HERE
raise NotImplementedError()
```

**Mandatory Task 1: The first million digits of $\pi$**

Read in the contents of the text file `pi.txt` (in the same folder as this notebook; it contains the first million digits of $\pi$) into the string variable `my_pi`. **Remove all line breaks.** 

Then, implement the function `birthday_in_pi(birthday, pidigits=my_pi)` that takes 2 inputs: 
* birthday: `str`, in format `DDMMYY`, and
* pidigits: `str`

and returns either `True` or `False`, depending on whether the birthday (the six-digit string) could be found within the pi digits.

Test your function with `my_pi` (the text you read in from the file) and your own (or any other) birthday.

<!-- grade_id: cell-90f071a114f2e3bf, points: 0, grade: false, solution: true -->
```python
# Read in pi.txt and remove all linebreaks ("\n")
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9932ba23b7d334d2, points: 0, grade: false, solution: true -->
```python
# function that returns True/False depending on 
# whether birthday is contained in pidigits:
def birthday_in_pi(birthday, pidigits=my_pi):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-393a475780070499, points: 0, grade: false, solution: true -->
```python
# test the function with your own (or any other's) birthday
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-86a1d57624604fe6, points: 1, grade: true, solution: false -->
```python
# run this cell to check whether: 
# - you read in the file correctly
# - the function works correctly

assert type(my_pi)==str
assert my_pi[0:4]=="3.14"
assert my_pi[-6:-1]=="45815"
assert "\n" not in my_pi

assert birthday_in_pi("010170")==True
assert birthday_in_pi("241295", my_pi)==False
assert birthday_in_pi("080399", pidigits=my_pi)==True
```

**Mandatory Task 2: Data wrangling with pandas**

In the same folder as this notebook, you will find `air_quality_no2.csv`, a data set of a month's worth of hourly measurements of NO2 (nitrogen dioxide, an atmospheric pollutant) from measuring stations in Paris, Londn, and Antwerp. (Data set taken from [here](https://github.com/pandas-dev/pandas/blob/main/doc/data/air_quality_no2.csv))

Complete the following tasks:
* Read in the data set into a variable called `df`, using the `pandas` library 
* Remove the column `station_antwerp` from the dataframe
* Remove all rows that contain missing data in **any** of the two columns `station_london` or `station_paris` (or in both) from the data frame
* For how many of the hourly measurements Paris has a higher NO2 concentration than London? Save the answer (an integer number) to the variable `paris_higher`
* On what dates and times of day did the London and Paris stations have exactly the same NO2 levels? Using `.copy()`, create a new dataframe called `same_levels`, with only those rows from df where `station_paris` and `station_london` have the same values

<!-- grade_id: cell-ffe959e9564a883e, points: 0, grade: false, solution: true -->
```python
# Solve task 2

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-44f029ef214851b1, points: 1, grade: true, solution: false -->
```python
# make sure that the antwerp column has been dropped from the df
assert "station_antwerp" not in df.columns

# make sure that all rows with missing values have been removed
assert sum(df["station_paris"].isna())==0
assert sum(df["station_london"].isna())==0

# make sure that the number of measurements where Paris > London has been computed correctly
assert paris_higher == 472

# make sure that the day with same NO2 levels for Paris and London has been found 
assert samelevels.loc[:, "datetime"].values == "2019-06-07 15:00:00"
```

**Mandatory Task 3** 

Write the function `recursive_sum` that takes an integer `n` as input, and computes the sum of all integers up to **and including** n as output, **recursively**. 

That means that the function 
* should return `n` in the base case (if n<=1)
* should **call itself in the return statement** in all other cases (if n>1), in a way that allows to compute the sum of all integers up to and including `n`

<!-- grade_id: cell-76caa9732c889607, points: 0, grade: false, solution: true -->
```python
# Solve Task 3
def recursive_sum(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-af8ac661615e55c0, points: 1, grade: true, solution: false -->
```python
# check that the function computes the sum of all integers up to n correctly
assert recursive_sum(5)==sum([1,2,3,4,5])
assert recursive_sum(8)==sum([1,2,3,4,5,6,7,8])
assert recursive_sum(20)==sum([i for i in range(1,21)])

# check that in the base case (n=1 or n=0) the function returns the correct value of n:
assert recursive_sum(1)==1
assert recursive_sum(0)==0
```

**Mandatory Task 4: from HTTP request to text file**

> Note: The `request` module is presented in lecture 10. 

Under the link below, you will find an extract from the FAQs of the Python documentation, explaining why it is called "Python". Your task:
* Send an HTTP request to the link below (use the built-in `requests` module); 
* save the text attribute of the HTTP response into the string variable `faq`; 
* then, save the string as a text file, `faq.txt` (in the same folder as this notebook).

Link: https://raw.githubusercontent.com/anastassiavybornova/pythoncrashcourse/main/assignments/assignment04/whyisitcalledpython.txt

<!-- grade_id: cell-1c8f7ee7f1542dce, points: 0, grade: false, solution: true -->
```python
# Solve Task 4
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-166276a50bff8aa8, points: 1, grade: true, solution: false -->
```python
# run this cell to check if you 
# - read in the text with the request module correctly
# - wrote the text to a file correcly

# text was imported correctly
assert type(faq)==str
assert len(faq)==475
assert faq[128:156]=="Monty Python’s Flying Circus"

# file exists & contains the same thing as the variable "sentence"
assert "faq.txt" in os.listdir()
assert open("faq.txt", "r").read() == faq
```

**Mandatory Task 5** 

Explain, **in your own words**, how the binary search algorithm works. You can use the ["guessing game" from lecture09](https://learnit.itu.dk/mod/resource/view.php?id=190513) (guessing a number from 1 to 1000) as an example.

(150 words max)

<!-- grade_id: cell-15e5bdd522cf94a7, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE