# Assignment 02

for the course [Introduction to Data Science and Programming](https://learnit.itu.dk/course/view.php?id=3022199) (Autumn 2023) at ITU

**Due: September 25, 23:59**

Before you turn this assignment in, make sure everything runs as expected. (Go to `Restart and run all` - should not throw any assertion errors.) 

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and ITU email address below:

```python
NAME = ""
ITU_MAIL = "" # xxxx@itu.dk
```

***
**Mandatory Task 1:** Write the function `sum_list` that takes a list as input and returns the sum of all elements in the list as ouput. Do **not** use Python's built-in `sum()` function for this task; the task is to write this function from scratch. It is fine if your function throws an error if anything else than a list of integers is entered as input.

<!-- grade_id: cell-c3b8fc79af83886c, points: 0, grade: false, solution: true -->
```python
# Solve Task 1
def sum_list(my_list):
    '''
    takes a list as input and returns the sum of all elements in the list as output
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

Make sure your function passes all the tests: running the cell below should not throw any errors! If it does, it means you need to adjust your function. You can run this cell as often as you need to, to make sure your function works.

<!-- grade_id: cell-35af3afdbf1678c3, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 1 correctly by running this cell
assert sum_list([1,2,3])==6
assert sum_list([9])==9
assert sum_list([0,0,-1,-10])==-11

```

***
**Mandatory Task 2:** Write the function `is_even` that takes an integer as input and returns either `True` (if the input integer is even), `False` (if the input integer is odd), **or** the string `"wrong input"` (if the input is not an integer).

<!-- grade_id: cell-a3ec09c0ee06a3e1, points: 0, grade: false, solution: true -->
```python
# Solve Task 2
def is_even(n):
    '''
    takes an integer as input and returns either 
    True (if the input integer is even), 
    False (if the input integer is odd), **or** the string 
    "wrong input" (if the input is not an integer).
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7905e96f23995be4, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 2 correctly

# we correctly detect even numbers
assert is_even(2)==True
assert is_even(16)==True
assert is_even(0)==True

# we correctly detect odd numbers
assert is_even(1)==False
assert is_even(7659)==False
assert is_even(-3)==False

# we correctly detect non-integer input
assert is_even(7.5)=="wrong input"
assert is_even("inputstring")=="wrong input"
assert is_even([1,2,3])=="wrong input"
```

***
**Mandatory Task 3:** Write the function `compute_factorial` that takes an integer $n>=1$ as input and returns the factorial $n!$ as output. If the input n is anything else than a positive integer, the function should return the number $0$. Do **not** use any built-in functions or modules in Python for this task. **Hint:** use a `for` loop with a `range()` to compute the factorial. 

<!-- grade_id: cell-e8db086fc0ce7341, points: 0, grade: false, solution: true -->
```python
# Solve Task 3
def compute_factorial(n):
    '''
     takes an integer n>=1 as input and returns the factorial n! as output.
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-35c12a7e2e2818f4, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 3 correctly

# Computing factorials correctly
assert compute_factorial(1)==1
assert compute_factorial(3)==6
assert compute_factorial(10)==3628800

# Returning 0 if number is not an integer >= 1
assert compute_factorial(0)==0
assert compute_factorial(-7)==0
assert compute_factorial("thirty-three")==0
```

***
**Mandatory Task 4** Write a function `find_largest_factorial_below` that takes a positive number $n$ (`int` or `float`) as input and returns the **largest possible** number $a$ for which $a! \leq n$. In other words, this function returns the number $a$ whose factorial $a!$ gets as close as possible to $n$ without growing larger than $n$. **Hint:** Use a `while` loop with within your function.

<!-- grade_id: cell-8c5e7d0a48ac173b, points: 0, grade: false, solution: true -->
```python
# Solve Task 4
def find_largest_factorial_below(n):
    '''
    takes a positive number n (int or float) as input and 
    returns the largest possible number a for which a! <= n
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d583b2ef04bcd208, points: 1, grade: true, solution: false -->
```python
# Check if you solved Task 4 correctly

# For numbers between 0 and 2, it should return 1
assert find_largest_factorial_below(1.5)==1

# For numbers between 6 and 24, it should return 3
assert find_largest_factorial_below(17)==3

# For 24, it should return 4
assert find_largest_factorial_below(24)==4

# For numbers <= 0 and non-numeric input, it should return 0
assert find_largest_factorial_below(0)==0
assert find_largest_factorial_below(-3)==0
assert find_largest_factorial_below("mystring")==0

```

***
**Mandatory Task 5:** **In your own words**, explain the main similarities and differences between `for` and `while` loops in Python. Use the two functions `compute_factorial` and `find_largest_factorial_below` from Tasks 3 and 4 as examples. 

**150 words max** 


<!-- grade_id: cell-7e11d193408739d6, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE